---
title: Pengenalan
description: 'Chapter description goes here.'
---

## Fungsi print

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

Pada bagian ini Kamu akan belajar mengenai dasar-dasar pemrograman dengan python

`@instructions`
Pada bagian ini, Kamu dapat menuliskan nama Kamu pada bagian editor

`@hint`
Tuliskan nama kamu pada bagian editor

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
